<a href="https://colab.research.google.com/github/Tabeaz/cs-tutoring/blob/master/Assignment_0.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Assignment 0: Programming Basics

## Introduction: A Simple Calculation Operator as your first computer program!
This assignment will introduce the following arithmetic operations:
* Addition: `+`
* Subtraction: `-`
* Multiplication: `*`
* Division: `/`
* Power Function: `**`

## Setup

First, we need two variables that we want to add:

Define the variables `x` and `y` and assign them the integer values `10` (for `x`) and `4` (for `y`):

```python
# Add your code in this cell:
x = 10
y = 4

print(x)
print(y)


```

Output:
```
10
4
```

### Testing?
How can you test if your code is working correctly? Test your own code:

Check if you defined the variable `x` by `print`-ing it:
```python
print(x)
```
If you see the value `10` below the cell then it worked! If you you see an error, __read__ it and try to understand what it means. __Do not ignore__ any error or be scared of the rather long output. Check where the error occurs: _What does the error say?_, _What did Python expect but what was actually given?_


In these assignments, we will use *automated unit test* to verify the implementation. This means that the code below will call your functions (in later assignments), test if the variables you defined actually exist and if their value is correct.

<!-- grade_id: task_1.1, points: 1, grade: true, solution: false -->
```python
# This is a very simple test if you defined variable x. You do not need to understand it fully.
# If you get any error here, this means that your implementation above is incorrect.
# Read the error message!
assert locals().get('x'), 'You did not define variable x.'
assert isinstance(x, int), 'You did not define x to be an integer.'
assert x == 10, 'You did not define x to be 10.'

# This is the same test but for variable y.
assert locals().get('y'), 'You did not define variable y.'
assert isinstance(y, int), 'You did not define y to be an integer.'
assert y == 4, 'You did not define y to be 4.'


print('If you see this message printed, then the test succeeded.')
```

The next test fails __intentionally__. It will test if there is variable defined with the name `h`. Check the error message, it tells you that the variable `h` is not defined!

```python
assert locals().get('h'), 'You did not define variable h'
```

## Addition, Subtraction, and Multiplication

Now that we have the variable `x` and `y` in scope, we can do calculations.

* Define variable `z1` to be the result of adding the two _variables_ `x` and `y`.
* Define variable `z2` to be the result of subtracting the two _variables_ `x` and `y`.
* Define variable `z3` to be the result of multiplying the two _variables_ `x` and `y`.

```python
z1 = x + y
z2 = x - y
z3 = x * y

print(z1, z2, z3)


```

Output:
```
14 6 40
```

```python
# Again, here we test your implementation
assert locals().get('z1'), 'You did not define variable z1.'
assert z1 == 14, 'Incorrect result for z1.'

assert locals().get('z2'), 'You did not define variable z2.'
assert z2 == 6, 'Incorrect result for z2.'

assert locals().get('z3'), 'You did not define variable z3.'
assert z3 == 40, 'Incorrect result for z3.'

print('If you see this message printed, then the test succeeded.')
```

## Division

A division is a bit different. The result of a division is __always__ of type `float`, no matter if the two operands are `int`, `float`, or a mixture.
Additionally, a division is dangerous because a division by `0` is not possible.

Try the division of `x` by `0`.

```python
z4 = x / 0
print(z4)
```

Output:
```
ZeroDivisionError                         Traceback (most recent call last)
<ipython-input-4-cbd47e886e5a> in <module>()
----> 1 z4 = x / 0
      2 print(z4)

ZeroDivisionError: division by zero
```

See that there is a well-defined error message for this exact type or error? It is called `ZeroDivionError` and it occurs if there a divison by zero.

For this task, store the result of the division of `x` by `y` in the variable `z4`.

```python
z4 = x / y
print(z4)
```

Output:
```
2.5
```

```python
# Again, here we test your implementation
assert locals().get('z4'), 'You did not define variable z4.'
assert z4 == 2.5, 'Incorrect result for z4.'

print('If you see this message printed, then the test succeeded.')
```

## Combination

Now, we want to combine a few arithmetic operation:

1. Define the variables `x`, `y`, and `z` with the values `3`, `21`, and `12`.
2. Define a new variable `r` with the value `0`.
3. Add `x` to `r` (store the result in `r`).
4. Add the result of the multiplication of `y` and `z` to `r` (store the result in `r`).
5. Divide `r` by `x` (store the result in `r`).
6. Subtract the result of the substraction of `x` and `y` from `r`  (store the result in `r`).

```python
x = 3
y = 21
z = 12
r = 0

r = x + r
r = y * z + r
r = r / x
r = r - (x - y)

print(r)
```

Output:
```
103.0
```

```python
# The result is supposed to be 103.0
assert locals().get('r'), 'You did not define variable r.'
assert r == 103.0, 'Incorrect result for r.'

print('If you see this message printed, then the test succeeded.')
```

## Combination

Another example:

1. Define the variables `x`, `y`, and `z` with the values `2.4`, `4`, and `14`.
2. Define the variable `r` with the initial value `2`.
3. Add `x`, `y`, and `z` in one step and multiply the result with `r` (store in `r`).
4. Add the square of `y` to `r` (store in `r`).
5. Divide `r` by the result of the addition of `y` and `1` (store in `r`).
6. Multiply `r` by `y` to the power of the division of `y` by `2` (store in `r`).

```python
x = 2.4
y = 4
z = 14
r = 2

r = (x+y+z) * r
r = r + y**2
r = r / (y + 1)
r = r * y**(y/2)
print(r)
```

Output:
```
181.76
```

```python
assert locals().get('r'), 'You did not define variable r.'
assert r == 181.76, 'Incorrect result for r.'

print('If you see this message printed, then the test succeeded.')
```